For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

## Solve

4 numbers are multiplies of 8, sum 336?


Use brute force if possible!

Other solution can be:

    import sympy
    x = sympy.var('x')
    s = sympy.solve(8*x+8*(x+1)+8*(x+2)+8*(x+3)-368,x)[0]
    numbers = [8*(s+i) for i in range(4)]

<!-- grade_id: cell-f761c36a52f85bfe, points: 0, grade: false, solution: true -->
```python
numbers = []
### BEGIN SOLUTION
for s in range(20):
    numbers = [(i+s)*8 for i in range(4)]
    print(numbers,':::' ,sum(numbers))

    if sum(numbers) == 336:
        break
### END SOLUTION
```

Output:
```
[0, 8, 16, 24] ::: 48
[8, 16, 24, 32] ::: 80
[16, 24, 32, 40] ::: 112
[24, 32, 40, 48] ::: 144
[32, 40, 48, 56] ::: 176
[40, 48, 56, 64] ::: 208
[48, 56, 64, 72] ::: 240
[56, 64, 72, 80] ::: 272
[64, 72, 80, 88] ::: 304
[72, 80, 88, 96] ::: 336
```

<!-- grade_id: cell-015ea81bcf5f9ee0, points: 1, grade: true, solution: false -->
```python
assert sum(numbers) == 336 
assert len(numbers) == 4

for i in range(1,4):
    assert numbers[i]-numbers[i-1] == 8

```